<!-- grade_id: cell-272a42f4ce1df5f9, points: 0, grade: false, solution: true -->
```python
def fibonacci_generator():
    """Generates all Fibonacci numbers."""
    # YOUR CODE HERE
    f1, f2 = 0, 1
    while True:
      yield f1
      f1, f2 = f2, f1 + f2
```

```python
r = []
for n in fibonacci_generator():
    r.append(n)
    if n > 100:
        break
assert r == [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144]

```

<!-- grade_id: cell-1e461d53e390cba5, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
c = 0
for n in fibonacci_generator():
    if n > 100:
        c += 1
    if n in [32, 88, 99, 231, 1000]:
        print("bad:", n)
        raise Exception()
    if n > 1000:
        break
assert c == 6
### END HIDDEN TESTS

```

<!-- grade_id: cell-6486f8e9d1595cc6, points: 0, grade: false, solution: true -->
```python
def prime_number_generator():
    """This generator returns all prime numbers."""
    # YOUR CODE HERE
    num = 2
    prime = [2]
    while True:
      num += 1
      check = all(num % number != 0 for number in prime)
      if check == True:
        yield prime[-1]
        prime.append(num)
        num += 1
      else:
        num += 1

```

```python
initial_primes = [2, 3, 5, 7, 11, 13, 17, 19, 23]

idx = 0
for p in prime_number_generator():
    assert p == initial_primes[idx]
    idx += 1
    if idx == len(initial_primes):
        break

```

<!-- grade_id: cell-a7e50a3c136cec77, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
c = 0
for n in prime_number_generator():
    if n > 900:
        c += 1
    if n in [32, 88, 99, 231, 1000]:
        print("bad:", n)
        raise Exception()
    if n > 1000:
        break
assert c == 15
### END HIDDEN TESTS

```