```python
import matplotlib.pyplot as plt

def draw_polygon(*coords, closed=True):
    co = list(coords)
    fix, ax = plt.subplots(1, 1)
    if closed and co[0] != co[-1]:
        co.append(co[0]) # Repeat first point
    xs, ys = zip(*co) # Splits coordinates
    ax.set_aspect('equal', 'box')
    ax.plot(xs, ys)
    plt.show()

```

```python
draw_polygon((0, 0), (1, 2), (2, 0))

```

```python
draw_polygon((0, 0), (1, 0), closed=False)

```

```python
draw_polygon((0, 0), (0.25, 0), (0.5, 0.25), (0.75, 0), (1, 0), closed=False)

```

```python
draw_polygon((0, 0), (0, -1), closed=False)

```

```python
draw_polygon((0, 0), (0, -0.25), (0.25, -0.5), (0, -0.75), (0, -1), closed=False)

```

<!-- grade_id: cell-c8a16fb8e2951a5e, points: 0, grade: false, solution: true -->
```python
def kink(p0, p1):
    """Returns the tuple of the four points q0, q1, q2, q3, resulting from kinking p0, p1.
    Note that each q is returned as a two-element tuple."""
    # YOUR CODE HERE
    q0 = p0
    direction = (p1[0] - p0[0], p1[1] - p0[1])
    v = (-1 * direction[1], direction[0])
    v2 = (tuple(i * 0.25 for i in v))
    direction1 = (tuple(i * 0.25 for i in direction))
    direction2 = (tuple(i * 0.5 for i in direction))
    direction3 = (tuple(i * 0.75 for i in direction))
    q1 = tuple(a + b for a, b in zip(q0, direction1))
    q2 = tuple(a + b +c for a, b, c in zip(direction2, v2, q0))
    q3 = tuple(a + b for a, b in zip(q0, direction3))
    final = (q0, q1, q2, q3)
    return final
```

<!-- grade_id: cell-912ef8324b82d50b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def kink(p0, p1):
    """Returns the tuple of the four points q0, q1, q2, q3, resulting from kinking p0, p1.
    Note that each q is returned as a two-element tuple."""
    q0 = p0
    direction = (p1[0] - p0[0], p1[1] - p0[1])
    v = (-1 * direction[1], direction[0])
    v2 = (tuple(i * 0.25 for i in v))
    direction1 = (tuple(i * 0.25 for i in direction))
    direction2 = (tuple(i * 0.5 for i in direction))
    direction3 = (tuple(i * 0.75 for i in direction))
    q1 = tuple(a + b for a, b in zip(q0, direction1))
    q2 = tuple(a + b + c for a, b, c in zip(direction2, v2, q0))
    q3 = tuple(a + b for a, b in zip(q0, direction3))
    final = (q0, q1, q2, q3)
    return final

```

```python
assert kink((0, 0), (1, 0)) == ((0, 0), (0.25, 0.0), (0.5, 0.25), (0.75, 0.0))
assert kink((0, 0), (0, -1)) == ((0, 0), (0.0, -0.25), (0.25, -0.5), (0.0, -0.75))

```

```python
def approxeq(list1, list2, epsilon=0.001):
    assert len(list1) == len(list2)
    for p1, p2 in zip(list1, list2):
        x1, y1 = p1
        x2, y2 = p2
        assert abs(x2 - x1) < epsilon and abs(y2 - y1) < epsilon

```

```python
approxeq(kink((0, 0), (1, 0)), ((0, 0), (0.25, 0.0), (0.5, 0.25), (0.75, 0.0)))

```

<!-- grade_id: cell-73e1f43f0ff453dc, points: 10, grade: true, solution: false -->
```python
approxeq(kink((0, 0), (4, 5)), ((0, 0), (1.0, 1.25), (0.75, 3.5), (3.0, 3.75)))
approxeq(kink((4, 3), (-9, -1)), ((4, 3), (0.75, 2.0), (-1.5, -2.25), (-5.75, 0.0)))

### BEGIN HIDDEN TESTS
approxeq(kink((-3, -4), (3, -1)), ((-3, -4), (-1.5, -3.25), (-0.75, -1.0), (1.5, -1.75)))
approxeq(kink((3, 5), (1, 12)), ((3, 5), (2.5, 6.75), (0.25, 8.0), (1.5, 10.25)))
### END HIDDEN TESTS

```

```python
poly1 = [(0, 0), (1, 2), (2, 0)]
draw_polygon(*poly1)

```

<!-- grade_id: cell-2d5cc9db47362620, points: 0, grade: false, solution: true -->
```python
def kink_polygon(poly):
    # YOUR CODE HERE
  final = []
  for tup in range(0, (len(poly)-1)):
    kinks = kink(poly[tup], poly[tup+1])
    last = kink(poly[-1], poly[0])
    final += kinks
    if last not in final:
      final += last
  return final
```

<!-- grade_id: cell-6bb9ccf63259d739, points: 0, grade: false, solution: true -->
```python
# You can use this space to debug your code.

# YOUR CODE HERE
def kink_polygon(poly):
    # YOUR CODE HERE
  final = []
  for tup in range(0, (len(poly)-1)):
    kinks = kink(poly[tup], poly[tup+1])
    final += kinks
  last = kink(poly[-1], poly[0])
  final += last
  return final
```

```python
poly1 = [(0, 0), (1, 2), (2, 0)]
poly2 = kink_polygon(poly1)
draw_polygon(*poly2)

```

```python
poly3 = kink_polygon(poly2)
draw_polygon(*poly3)

```

```python
poly1 = [(0, 0), (1, 2), (2, 0)]
poly2 = kink_polygon(poly1)
approxeq(poly2, [(0, 0),
 (0.25, 0.5),
 (0.0, 1.25),
 (0.75, 1.5),
 (1, 2),
 (1.25, 1.5),
 (2.0, 1.25),
 (1.75, 0.5),
 (2, 0),
 (1.5, 0.0),
 (1.0, -0.5),
 (0.5, 0.0)])
poly3 = kink_polygon(poly2)
approxeq(poly3, [(0, 0),
 (0.0625, 0.125),
 (0.0, 0.3125),
 (0.1875, 0.375),
 (0.25, 0.5),
 (0.1875, 0.6875),
 (-0.0625, 0.8125),
 (0.0625, 1.0625),
 (0.0, 1.25),
 (0.1875, 1.3125),
 (0.3125, 1.5625),
 (0.5625, 1.4375),
 (0.75, 1.5),
 (0.8125, 1.625),
 (0.75, 1.8125),
 (0.9375, 1.875),
 (1, 2),
 (1.0625, 1.875),
 (1.25, 1.8125),
 (1.1875, 1.625),
 (1.25, 1.5),
 (1.4375, 1.4375),
 (1.6875, 1.5625),
 (1.8125, 1.3125),
 (2.0, 1.25),
 (1.9375, 1.0625),
 (2.0625, 0.8125),
 (1.8125, 0.6875),
 (1.75, 0.5),
 (1.8125, 0.375),
 (2.0, 0.3125),
 (1.9375, 0.125),
 (2, 0),
 (1.875, 0.0),
 (1.75, -0.125),
 (1.625, 0.0),
 (1.5, 0.0),
 (1.375, -0.125),
 (1.375, -0.375),
 (1.125, -0.375),
 (1.0, -0.5),
 (0.875, -0.375),
 (0.625, -0.375),
 (0.625, -0.125),
 (0.5, 0.0),
 (0.375, 0.0),
 (0.25, -0.125),
 (0.125, 0.0)])

```

<!-- grade_id: cell-c040a1dcdf1c8c1a, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
poly = ((1, 1), (2, 4), (4, 4), (5, 1))
approxeq(kink_polygon(poly), [(1, 1),
 (1.25, 1.75),
 (0.75, 2.75),
 (1.75, 3.25),
 (2, 4),
 (2.5, 4.0),
 (3.0, 4.5),
 (3.5, 4.0),
 (4, 4),
 (4.25, 3.25),
 (5.25, 2.75),
 (4.75, 1.75),
 (5, 1),
 (4.0, 1.0),
 (3.0, 0.0),
 (2.0, 1.0)])
### END HIDDEN TESTS

```

<!-- grade_id: cell-dff97bbe3d63824, points: 0, grade: false, solution: true -->
```python
def poly_area(polygon):
    """Returns the area of a polygon."""
    # YOUR CODE HERE
    area = 0
    for i in range(0, (len(polygon)-1)):
      x1 = polygon[i][0]
      x2 = polygon[i+1][0]
      y1 = polygon[i][1]
      y2 = polygon[i+1][1]
      first = (x2 - x1)
      second = (y1+y2)/2
      factorial = first * second
      area += factorial
    return area

```

<!-- grade_id: cell-714f84d8d049fe0a, points: 0, grade: false, solution: true -->
```python
import math

def poly_perimeter(polygon):
    """Returns the perimeter of a polygon."""
    # YOUR CODE HERE
    perimeter = 0
    for i in range(0, (len(polygon)-1)):
      x1 = polygon[i+1][0]
      y1 = polygon[i+1][1]
      first = ((x1)-polygon[i][0]) ** 2
      second = ((y1)-polygon[i][1]) ** 2
      peri = math.sqrt(first + second)
      perimeter += peri
    third = ((polygon[-1][0])-polygon[0][0]) ** 2
    fourth = ((polygon[-1][1])-polygon[0][1]) ** 2
    perim = math.sqrt(third + fourth)
    perimeter += perim
    return perimeter

```

<!-- grade_id: cell-61df2a5fd9bb45b0, points: 0, grade: false, solution: true -->
```python
## Here you can test your code.

# YOUR CODE HERE
def poly_area(polygon):
    """Returns the area of a polygon."""
    # YOUR CODE HERE
    area = 0
    for i in range(0, (len(polygon)-1)):
      x1 = polygon[i][0]
      x2 = polygon[i+1][0]
      y1 = polygon[i][1]
      y2 = polygon[i+1][1]
      first = (x2 - x1)
      second = (y1+y2)/2
      factorial = first * second
      area += factorial
    return area

import math

def poly_perimeter(polygon):
    """Returns the perimeter of a polygon."""
    # YOUR CODE HERE
    perimeter = 0
    for i in range(0, (len(polygon)-1)):
      x1 = polygon[i+1][0]
      y1 = polygon[i+1][1]
      first = ((x1)-polygon[i][0]) ** 2
      second = ((y1)-polygon[i][1]) ** 2
      peri = math.sqrt(first + second)
      perimeter += peri
    third = ((polygon[-1][0])-polygon[0][0]) ** 2
    fourth = ((polygon[-1][1])-polygon[0][1]) ** 2
    perim = math.sqrt(third + fourth)
    perimeter += perim
    return perimeter

```

<!-- grade_id: cell-f50aa72f7a51a6d0, points: 10, grade: true, solution: false -->
```python
poly1 = [(0, 0), (1, 2), (2, 0)]
assert poly_area(poly1) == 2
assert abs(poly_perimeter(poly1) - 6.472) < 0.01

### BEGIN HIDDEN TESTS
poly1 = [(0, 0), (0, 1), (1, 1), (1, 0)]
assert poly_area(poly1) == 1
assert abs(poly_perimeter(poly1) - 4) < 0.01
### END HIDDEN TESTS

```

```python
poly = [(0, 0), (1, 2), (2, 0)]
for n in range(6):
    poly = kink_polygon(poly)
    draw_polygon(*poly)

```

```python
poly = [(0, 0), (0, 1), (1, 1), (1, 0)]
for n in range(7):
    poly = kink_polygon(poly)
    area = poly_area(poly)
    perim = poly_perimeter(poly)
    print("Area:", area, "Perimeter:", perim)
assert 14.9 < perim < 15
assert 1.399 < area < 1.4

```

```python
poly = [(0, 1), (1, 4), (6, 4), (8, 1)]
for n in range(7):
    poly = kink_polygon(poly)
    area = poly_area(poly)
    perim = poly_perimeter(poly)
    print("Area:", area, "Perimeter:", perim)
assert 30.6 < area < 31
assert 73 < perim < 74
```