```python
import math

class Complex(object):

    def __init__(self, r, i):
        self.r = r # Real part
        self.i = i # Imaginary part

    def __add__(self, other):
        return Complex(self.r + other.r, self.i + other.i)

    def __sub__(self, other):
        return Complex(self.r - other.r, self.i - other.i)

    def __mul__(self, other):
        return Complex((self.r * other.r - self.i * other.i),
                       (self.r * other.i + self.i * other.r))

    @property
    def modulus_square(self):
        return self.r * self.r + self.i * self.i

    @property
    def modulus(self):
        return math.sqrt(self.modulus_square)

    def inverse(self):
        m = self.modulus_square # to cache it
        return Complex(self.r / m, - self.i / m)

    def __truediv__(self, other):
        return self * other.inverse()

    def __repr__(self):
        """This defines how to print a complex number."""
        if self.i < 0:
            return "{}-{}i".format(self.r, -self.i)
        return "{}+{}i".format(self.r, self.i)

    def __eq__(self, other):
        """We even define equality"""
        return self.r == other.r and self.i == other.i

```

<!-- grade_id: cell-796a797c822736bc, points: 0, grade: false, solution: true -->
```python
# Definition of Fraction class

# Here is the gcd function, as it may well be useful to you.
def gcd(m, n):
    # This is the "without loss of generality" part.
    m, n = (m, n) if m > n else (n, m)
    m, n = abs(m), abs(n)
    return m if n == 0 else gcd(m % n, n)

class Fraction(object):

    def __init__(self, numerator, denominator):
        assert isinstance(numerator, int)
        assert isinstance(denominator, int)
        assert denominator != 0
        # YOUR CODE HERE
        d = gcd(numerator, denominator)
        gcd(numerator, denominator)
        denominator1 = denominator / d
        numerator1 = numerator / d
        if denominator < 0:
          numerator1 = numerator1 * (-1)
          denominator1 = denominator1 * (-1)
        self.numerator = int(numerator1)
        self.denominator = int(denominator1)

    def __repr__(self):
        """Pretty print a fraction."""
        return "{}/{}".format(self.numerator, self.denominator)

    ## Here, implement the methods for +, -, *, /, =, and <.
    ## Done quite at leisure, with spaces and all, this can be done in about
    ## 25 lines of code.

     # YOUR CODE HERE
    def __add__(self, second):
      if self.denominator == second.denominator:
        add = second.numerator + self.numerator
        return Fraction(add, self.denominator)
      if self.denominator != second.denominator:
        num_self = self.numerator * second.denominator
        den = self.denominator * second.denominator
        num_second = second.numerator * self.denominator
        numer = num_self + num_second
        denom = den
        return Fraction(numer, denom)

    def __sub__(self, second):
      if self.denominator == second.denominator:
        sub = self.numerator - second.numerator
        return Fraction(sub, self.denominator)
      if self.denominator != second.denominator:
        num_self = self.numerator * second.denominator
        den = self.denominator * second.denominator
        num_second = second.numerator * self.denominator
        numer = num_self - num_second
        denom = den
        return Fraction(numer, denom)

    def __mul__(self, second):
      num = self.numerator * second.numerator
      den = self.denominator * second.denominator
      g = gcd(num, den)
      numer = int(num/g)
      denom = int(den/g)
      return Fraction(numer, denom)

    def __truediv__(self, second):
      num = self.numerator * second.denominator
      den = self.denominator * second.numerator
      g = gcd(num, den)
      numer = int(num/g)
      denom = int(den/g)
      return Fraction(numer, denom)

    def __eq__(self, second):
      if (second.numerator== self.numerator and second.denominator == self.denominator):
        return True
      if (second.numerator != self.numerator or second.denominator != self.denominator):
        return False
  
    def __lt__(self, second):
      if (self.numerator, self.denominator < second.numerator, second.denominator):
        return True
      if (second.numerator, second.denominator < self.numerator, self.denominator):
        return True

```

```python
## First, let us check that you correctly put the fraction into normal form,
## without common factor between numerator and denominator, and with a
## positive denominator.

f = Fraction(8, 6)
assert f.numerator == 4 and f.denominator == 3

f = Fraction(-8, 6)
assert f.numerator == -4 and f.denominator == 3

f = Fraction(8, -6)
assert f.numerator == -4 and f.denominator == 3

f = Fraction(-8, -6)
assert f.numerator == 4 and f.denominator == 3

f = Fraction(0, 10)
assert f.numerator == 0 and f.denominator == 1

```

<!-- grade_id: cell-2cabfda268f93ed3, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
f = Fraction(15, 10)
assert f.numerator == 3 and f.denominator == 2
f = Fraction(3, -6)
assert f.numerator == -1 and f.denominator == 2
### END HIDDEN TESTS

```

```python
f = Fraction(8, 6) + Fraction(25, 20)
assert f.numerator == 31 and f.denominator == 12
assert f == Fraction(31, 12)
assert f == Fraction(62, 24)

assert Fraction(6, 4) + Fraction(-8, 6) == Fraction(6, 4) - Fraction(8, 6)
assert not (Fraction(6, 4) + Fraction(-8, 6) == Fraction(6, 5) - Fraction(8, 6))

```

<!-- grade_id: cell-a79483c948d8dea7, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert Fraction(3, 2) + Fraction(3, 2) == Fraction(3, 1)
assert Fraction(30, 25) - Fraction(8, 7) + Fraction(8, 7) == Fraction(30, 25)
assert not Fraction(4, 5) == Fraction(4, -5)
assert not Fraction(4, 5) == Fraction(0, 5)
### END HIDDEN TESTS

```

```python
assert Fraction(3, 2) * Fraction(2, 3) == Fraction(1, 1)
assert Fraction(3, 2) / Fraction(2, 3) == Fraction(9, 4)
assert Fraction(3, 2) / Fraction(6, 4) == Fraction(1, 1)
assert Fraction(32, 16) == Fraction(2, 1)
assert not Fraction(33, 16) == Fraction(4, 2)

```

<!-- grade_id: cell-1a6123cf3a7681e7, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert Fraction(12, 6) * Fraction(2, 3) == Fraction(4, 3)
assert Fraction(-20, 2) / Fraction(15, 3) == Fraction(-2, 1)
assert Fraction(3, 2) / Fraction(-4, 3) == Fraction(-9, 8)
assert not Fraction(2, -1) == Fraction(2, 1)
### END HIDDEN TESTS

```

```python
assert Fraction(5, 7) < Fraction(5, 6)
assert Fraction(-3, 2) < Fraction(0, 3)

```

<!-- grade_id: cell-37b6f3021bc2e4f, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert Fraction(6, 8) < Fraction(7, 8)
assert Fraction(-7, 8) < Fraction(6, 8)
### END HIDDEN TESTS

```

```python
a = Fraction(7, 8)
b = Fraction(-4, 5)
a + b
a / b
a < b
a * b
assert a == Fraction(7, 8)
assert b == Fraction(-4, 5)

```

```python
import random
for _ in range(1000):
    a = Fraction(random.randint(-200, 200), random.randint(1, 100))
    b = Fraction(random.randint(-200, 200), random.randint(1, 100))
    c = Fraction(random.randint(-200, 200), random.randint(1, 100))
    assert Fraction(-1, 1000) < (a - b) * (a - b)
    assert (a - b) * (a + b) == a * a - b * b
    z = Fraction(0, 1) # Zero, as a fraction.
    if not ((a == z) or (b == z) or (c == z)):
        assert (a / b) * b == (a / c) * c
        assert (a / b) * (a / c) == (a * a) / (b * c)
        assert (a / b) / (b / c) == (a * c) / (b * b)
        assert (a * a * b * c) / (a * c) == a * b

```

<!-- grade_id: cell-21aa45adc9bbc8fe, points: 0, grade: false, solution: true -->
```python
class Int(Fraction):

    # YOUR CODE HERE
    def __init__(self, integer):
      super().__init__(integer, 1)

```

```python
assert Int(3) / Int(2) == Fraction(3, 2)
assert Int(3) * Int(4) / (Int(5) + Int(2)) == Fraction(12, 7)
assert Int(3) * Int(4) / (Int(5) + Int(1)) == Fraction(2, 1)

```